text_house = ["NEED TO GET FOOD",
              "CANNOT GO BACK YET",
              "I MUST BE STRONG",
              "I AM NOT DONE",
              "I STILL HAVE THINGS TO DO",
              "MUST GET BREAD",
              "MAN HAS NEEDS",
              "I SHALL NOT COWER IN SHAME",
              "I AM NOT A POOR FELLOW",
              "BREAD... FIND BREAD...",
              "CANNOT SLEEP WITH AN EMPTY STOMACH",
              "HUNGRY... I WILL NOT GO HUNGRY...",
              "NOT YET OUT OF THE FIGHT...",
              "I PROMISED TO BE STRONG",
              "CANNOT COWER IN SHAME",
              "NOT BACK INTO MY CAVE JUST YET",
              "I CAN ALWAYS WAIT FOR THE ANGEL OF DEATH LATER",
              "THE COMFORT ZONE IS AN ILLUSION..."]

text_lightning = ["THE CITY GROWS LARGER",
                  "RUNNING OUT OF SPACE",
                  "MEMORIES MAY SHATTER",
                  "CANNOT HIDE FROM THE FUTURE",
                  "CHANGE IS NOT ALWAYS FOR THE BETTER",
                  "THIS MIGHT BE MY LAST STAND",
                  "HEARTS HAVE INERTIA",
                  "SOON I MAY BE POWERLESS",
                  "AN UNBEARABLE PRESSURE",
                  "STARS HAVE FADED OUT",
                  "THE SKY IS NO MORE",
                  "PENETRATING EYES",
                  "ALWAYS WATCHING",
                  "TOO MANY OF THEM",
                  "OH THE HUMANITY"]

text_bus_lightning = ["FORGET ABOUT THE HOUSE",
                      "THE DECISION HAS BEEN MADE",
                      "DO NOT LISTEN TO THE HORRORS",
                      "THE CITY WILL FORGET ME",
                      "MY TIME HERE IS OVER ANYWAY",
                      "IT WAS A GOOD HOUSE",
                      "MEMORIES... WHAT GOOD ARE MEMORIES ANYWAY",
                      "SOMETIMES THE PAST IS JUST THAT... THE PAST",
                      "I HAVE TO LET GO...",
                      "A MOST DIFFICULT PATH I HAVE CHOSEN",
                      "IT WAS JUST A HOUSE...",
                      "SMILE BECAUSE IT HAPPENED... SMILE...",
                      "GOODBYE, THE CITY AND THE HOUSE",
                      "LOOK AWAY, LOOK AWAY FROM THE HOUSE",
                      "WHO KNOWS, MAYBE I AM STILL YOUNG AFTER ALL"]

text_house_lightning = ["I HAD A GOOD LIFE",
                        "IT IS TIME TO GO",
                        "I REGRET NOTHING",
                        "PEACE, AT LAST",
                        "I SHALL GO TO SLEEP",
                        "NOTHING MATTERS ANYMORE",
                        "THE FINAL HOUR...",
                        "...HAPPILY EVER AFTER..."]

text_ending_bus = """
I feel like it has been a week, but the bus finally arrived on destination
late at night.

My legs hurt from sitting still all day, and my head aches from the continuous
roar of the engine. Not to mention the vibration... I have gone completely
numb.

Yet I am here, I stand on my feet, in the outskirts of a town never seen
before. Maybe the folks will be weirded out, for I am a complete stranger.
But I am sure we will get along eventually. I am an old man, I don't have
much to ask for.

I just have to stop for a moment and light a fire. I have memories to burn.
Memories that would forever haunt me otherwise...
"""

text_ending_house = """
The city... and the house...

The moment I am gone, the city will devour the house. Well, no matter...
The tables have turned, and now the ship goes down with the captain.

I do not regret anything. Such is the nature of civilization. Had it not
been the city, the nature would eventually reclaim all I took anyway.
I am just an old man who can not let go of the past... The comfort of the
familiar...

Now I return to the hands of the nature. Maybe I will find that to be
even more familiar.

Yes, indeed... Soon enough, the house will perish in the hands of the city.
I can not imagine in the hands of what the city will perish.
But surely, it too will.

Goodbye, the city and the house.
"""

text_ending_walkout = """
My legs just gave up like that, huh. Well, I AM an old man...
Where am I going? It is just nothingness all the way to the horizon.

Where did I think I was going? Was I just running from the city?
Did I think reaching for an immense void would save me from the
growing city?

In a way... In a way, maybe it did.

I will perish here, in my last moment. And the house... Without me,
the house too will perish soon enough. The city will claim the house,
and then... and then...

I resent. I regret. It should not have been this way!

I will haunt this city. I swear, I will make my return, one day! And
that day, no one will be spared. I will reclaim what should have been
eternally mine!

I promise. I will be back.
"""

text_ending_quit = """
My poor heart... could not cope with the growing city, and all that it
brought. My friends disappeared one by one, and my family...

The town grew into a city, and the houses became tall blocks of metal
and concrete. The grassland perished, and became a land of barren rock
and asphalt. If it had not been for me, the house too would surely be...

But... but I did give up! Maybe the city is stronger than me after all...

I am weak, and I do not deserve the comfort of the house which I could
not protect.

Maybe death will eventually grant me some peace.
"""

